---
title_meta  : Chapter 1
title       : Chapter on DataCamp
description : This is what a chapter file looks like on DataCamp

--- type:VideoExercise lang:r xp:50 skills:1
## First Video Exercise


*** =video_link
//player.vimeo.com/video/108225030

--- type:MultipleChoiceExercise lang:r xp:50 skills:1
## First Multiple Choice Exercise

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
hint comes here

*** =pre_exercise_code
```{r}
# pec
```

*** =sct
```{r}
test_mc(2) # if 2 is the correct option.
```

--- type:NormalExercise lang:r xp:100 skills:1
## First Normal Exercise

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{r}
# pec
```

*** =sample_code
```{r}
# sample code
```

*** =solution
```{r}
# solution code
```

*** =sct
```{r}
# sct code
```
